<problem display_name="Exercise 2" markdown="null">
  <text>
    <p>
      <b>ESTIMATED TIME TO COMPLETE: 7 minutes</b>
      <br/>
      <b>      Note that you will have to answer all questions before you can click the Check button.</b>
    </p>
    <script type="text/python" system_path="python_lib">
def test_str(expect, ans):
  # Strip away ' ' or " " wrapping string input answer
  # (check which, if any, are present - don't want to pull valid
  #  characters out of the center of a string, eg "That's crazy!",
  #  nor do we wish to accept an input with mismatched modifiers, 
  #  eg 'abc")

  if ans[0] == "'" and ans[-1] == "'": ans = ans.strip("'")
  elif ans[0] == '"' and ans[-1] == '"': ans = ans.strip('"')

  if expect[0] == "'": expect = expect.strip("'")
  elif expect[0] == '"': expect = expect.strip('"')

  return expect == ans
</script>
    <p>
For each of the expressions below, specify its type and value. If it generates an error, select type 'NoneType' and write the word 'error' (note this is a word, not a string, no quotes) in the box for the value. While you could simply type these expressions into your IDE, we encourage you to answer them directly since this will help reinforce your understand of basic Python expressions. </p>
    <p>
Assume we've made the following assignments: </p>
    <pre>
&gt; str1 = 'hello'
&gt; str2 = ','
&gt; str3 = 'world'
</pre>
    <p>
      <section class="hints">
        <div class="collapsible">
          <header>
            <a href="#" id="id71">Note: The Python 'in' operator</a>
          </header>
          <section id="id71">
            <p>The operators <code>in</code> and <code>not in</code> test for collection membership (a 'collection' refers to a string, list, tuple or dictionary - don't worry, we will cover lists, tuples and dictionaries soon!). The expression </p>
            <pre>
element in coll
</pre>
            <p>
 evaluates to <code>True</code> if <code>element</code> is a member of the collection <code>coll</code>, and <code>False</code> otherwise. </p>
            <p>
The expression </p>
            <pre>
element not in coll
</pre>
            <p>
 evaluates to <code>True</code> if <code>element</code> is <b class="bf">not</b> a member of the collection <code>coll</code>, and <code>False</code> otherwise. </p>
            <p>
Note this returns the negation of <code>element in coll</code> - that is, the expression <code>element not in coll</code>        is equivalent to the expression <code>not (element in coll)</code>. </p>
          </section>
        </div>
      </section>
    </p>
    <ol class="enumerate">
      <li>
        <p>
          <code>str1</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'hello'">
            <textline size="50" correct_answer="'hello'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1[0]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'h'">
            <textline size="50" correct_answer="'h'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1[1]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'e'">
            <textline size="50" correct_answer="'e'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1[-1]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'o'">
            <textline size="50" correct_answer="'o'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>len(str1)</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="int"/>
          </optionresponse>
          <stringresponse answer="5">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1[len(str1)]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="NoneType"/>
          </optionresponse>
          <stringresponse answer="error">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1 + str2 + str3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'hello,world'">
            <textline size="50" correct_answer="'hello,world'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1 + str2 + ' ' + str3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'hello, world'">
            <textline size="50" correct_answer="'hello, world'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str3 * 3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'worldworldworld'">
            <textline size="50" correct_answer="'worldworldworld'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>'hello' == str1</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="boolean"/>
          </optionresponse>
          <stringresponse answer="True">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>'HELLO' == str1</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="boolean"/>
          </optionresponse>
          <stringresponse answer="False">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>'a' in str3</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="boolean"/>
          </optionresponse>
          <stringresponse answer="False">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <pre>
str4 = str1 + str3
'low' in str4
					</pre>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="boolean"/>
          </optionresponse>
          <stringresponse answer="True">
            <textline size="50"/>
          </stringresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str3[1:3]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'or'">
            <textline size="50" correct_answer="'or'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str3[:3]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'wor'">
            <textline size="50" correct_answer="'wor'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str3[:-1]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'worl'">
            <textline size="50" correct_answer="'worl'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str1[1:]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'ello'">
            <textline size="50" correct_answer="'ello'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str4[1:9]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'elloworl'">
            <textline size="50" correct_answer="'elloworl'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str4[1:9:2]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'elwr'">
            <textline size="50" correct_answer="'elwr'"/>
          </customresponse>
        </p>
      </li>
      <li>
        <p>
          <code>str4[::-1]</code>
          <optionresponse>
            <optioninput options="('NoneType','int','float','boolean','string')" correct="string"/>
          </optionresponse>
          <customresponse cfn="test_str" expect="'dlrowolleh'">
            <textline size="50" correct_answer="'dlrowolleh'"/>
          </customresponse>
        </p>
      </li>
    </ol>
  </text>
</problem>
